---
format: 
  live-html:
    toc: true

execute:
  echo: true
  warning: false
  message: false

embed-resources: true

filters:
  - custom-numbered-blocks

custom-numbered-blocks:
  classes:
    Assumption:
      colors: [FFCDD2, F44336]
      boxstyle: foldbox.simple
      collapse: false
    Example:
      colors: [BBDEFB, 2196F3]
      boxstyle: foldbox.simple
      collapse: false
    Exercise:
      colors: [C8E6C9, 4CAF50]
      boxstyle: foldbox.simple
      collapse: false
    Technical-point:
      colors: [C8E6C9, 4CAF50]
      boxstyle: foldbox.simple
      collapse: true
---
{{< include ./_extensions/r-wasm/live/_knitr.qmd >}}
{{< include ./_extensions/r-wasm/live/_gradethis.qmd >}}

```{webr}
#| include: false
library(tidyverse)
set_theme(theme_bw())
```

```{r}
#| include: false
library(tidyverse)
n <- 20
linearData <- tibble(
  x=rnorm(n,0,4),
  y=0.5 * x - 3 + rnorm(n,0, 2)
  )
ojs_define(linearData)
```

# Simple Linear Regression
##  Fitting models to data {.unnumbered}
in the previous section, you were given a linear model - we knew the
values of $\beta_0$, $\beta_1$ and even $\sigma$. However, in most
contexts we don't know the true relationship between $X$ and $Y$ in
advance. Instead, we are given *data* and try to *infer* the values of
$\beta$ and $\sigma$ by analysing the data.

Indeed, at the end of the last session we generated data from a linear model. In preparation for this chapter, I simulated data from a similar such linear model - your task here (and in regression generally) is to try an make a (educated) guess about what linear model (summarised by it's values for $\beta$ and $\sigma$) produced this 
data. 

- [ ] Continue

## Estimating parameters

### Hats
lets distinguish our estimates from the 'true', underlying, parameter values by giving them $\hat{\text{hats}}$. so, we have

$$
\hat{\beta_0}: \text{Estimated intercept}
$$
$$
\hat{\beta_1}: \text{Estimated slope}
$$
$$
\hat{\sigma}: \text{Estimated standard deviation}.
$$

We also define $\hat{Y}$ as the predicted value of $Y$ given our estimated model:

$$
\hat{Y}=\hat{\beta_0}+\hat{\beta_1}X
$$
Looking at some data, we try to find the values of our parameters that best 'fit' its distribution. Adjust the sliders below to find a line that you think fits the data:

```{ojs}
//| panel: sidebar
//| echo: false
viewof b0adj_1 = Inputs.range([-5, 5], {step: 1, label: html`${tex`\hat{\beta_0}`}: Estimated Intercept (asjust)`})
viewof b1_1 = Inputs.range([-2, 2], {step: 0.01, label:  html`${tex`\hat{\beta_1}`}: Estimated Slope`})

b0_1 = (-3 - 0.5*b1_1) + b0adj_1

tex.block`\hat{Y} = ${b0_1} + ${b1_1}x`
```

```{ojs}
//| panel: fill
//| echo: false

xRange = d3.extent(linearData.x)

lineData_1 = xRange.map(x => ({x, y: b1_1 * x + b0_1}))

Plot.plot({
  marks: [
    Plot.line(lineData_1, { x: "x", y: "y" }),

    Plot.dot(transpose(linearData), { x: "x", y: "y", r: 3 })
  ],
  x: { domain: xRange },
  y: { domain: [-10, 5], label: "y" },
})
```

::: Exercise
What is your best estimate for the values of $\beta_0$ and $\beta_1$ that best fit the given data?

```{webr}
#| exercise: ex_2.1
#| envir: Ex2
#| define:
#|   - beta_0_hat
#|   - beta_1_hat
beta_0_hat <- ______
beta_1_hat <- ______
```
:::: {.solution exercise="ex_1"}
```{webr}
#| exercise: ex_2.1
#| solution: true
```
::::
```{webr}
#| exercise: ex_2.1
#| envir: Ex2
#| check: true
#| class: wait
grade_this({
  if (is.numeric(get("beta_0_hat", envir = .envir_result)) && is.numeric(get("beta_1_hat", envir = .envir_result))){pass("Good guess!")}
  fail("both \\$\\beta_1\\$ and beta_0 should be numbers")
})
```
:::

###

Ok great, so we have an estimate for our line-of-best-fit:
```{ojs}
//| echo: false
tex.block`\hat{Y} = ${beta_0_hat} + ${beta_1_hat}\cdot X`
```
but what do we mean by 'best' fit here? How do we evaluate this estimate and maybe compare it to other estimates?

-   [ ] Continue

### Residuals

Residuals are the difference between observed values of $Y$, and the value predicted by our estimated linear predictor $\hat{Y}$. We denote residuals with the letter $e$:
$$
e=Y-\hat{Y} = Y - \hat{\beta_0}+\hat{\beta_1}X. 
$$

Residuals, $e$, are similar to the errors, $\varepsilon$, that we encountered in chapter 1 - but they are distinct. In many situations, we do not know the 'acutual' values of $\beta_0$ an $\beta_1$ - so we cannot calculate the errors, $\varepsilon= Y- E[Y]= Y-\beta_0 + \beta_1 X$. However, we do  have our estimates, $\hat{\beta_0}$ and $\hat{\beta_1}$, so we can calculate residuals. 

Indeed - calculating residuals gives us a way to assess how well our estimated model fits the data.  We can think of the residuals as being the 'mismatch' between our estimated linear predictor and each data point. By *minimizing* the size of residuals (minimisiong the mismatch of our model to the data), we can get a better fit of our line to data.

- [ ] Continue 

### Estimating $\beta$ by minimising residuals

#### how low can the (squared) residuals go?

Below, the plot now also displays residuals for each data point.
Underneath the model equation is the *Sum of Squared Residuals* (SSR), which gives a measure of the absolute difference between our linear predictor and the observed outcomes. 

$$
SSR=\sum_{i=1}^{n} e_i^2
$$

This gives us a numerical measure of how well the line fits the data - see how low you can get the SSR by adjusting slope and intercept.

$$
e_i = Y_i - \hat{Y}_i
$$

```{ojs}
//| panel: sidebar
//| echo: false
viewof b0adj = Inputs.range([-5, 5], {step: 1, label: html`${tex`\hat{\beta_0}`}: Estimated Intercept`})
viewof b1 = Inputs.range([-2, 2], {step: 0.01, label:  html`${tex`\hat{\beta_1}`}: Estimated Slope`})

b0 = -3 + b0adj

residualData = transpose(linearData).map(d => {
  const yhat = b1 * d.x + b0;
  const res  = d.y - yhat;
  return {
    ...d,
    yhat,
    residual: res,
    sign: res >= 0 ? "positive" : "negative",
    absRes: Math.abs(res)
  };
})
  
SSRes = d3.sum(residualData, d => d.residual ** 2)

tex.block`\hat{Y} = ${b1}X + ${b0}`

tex.block`\sum_{i}r_i^{2} = ${SSRes.toFixed(2)}`
```

```{ojs}
//| panel: fill
//| echo: false
lineData = xRange.map(x => ({x, y: b1 * x + b0}))

Plot.plot({
  marks: [
    Plot.ruleX(residualData, {x: "x",
      y1: "y",
      y2: "yhat",
      stroke: "sign", 
      strokeOpacity: 0.75,
      strokeWidth: d => 1 + d.absRes}),

    Plot.line(lineData, { x: "x", y: "y" }),

    Plot.dot(transpose(linearData), { x: "x", y: "y", r: 3 })
  ],
  x: { domain: xRange },
  y: { domain: [-10,5], label: "y" },
})
```

::: Exercise
Based on this visualisation, what do you think is the minimum possible $SSR$ achievable?
```{webr}
#| exercise: Ex_2.2
#| envir: Ex2
#| define:
#|   - SSR
SSR <- ________
```
:::: {.solution exercise="Ex_2.2"}
```{webr}
#| exercise: Ex_2.2
#| solution: true
```
::::
```{webr}
#| exercise: Ex_2.2
#| check: true
#| class: wait
grade_this({
  if (is.numeric(get("SSR", envir = .envir_result))){
    pass("Good guess!")}
  fail("SSR should be a number!")
})
```
If you want you can update your $\hat{\beta}$ estimates too (otherwise just proceed):

```{webr}
#| exercise: ex_2.3
#| envir: Ex2
beta_0_hat <- ______

beta_1_hat <- ______
```

:::
::: Technical-point
#### Why squares?
The least squares estimator chooses $\hat \beta_0$ and $\hat \beta_1$ so that the sum of squared residuals $\sum_{i=1}^{n} e_i^2$ is as small as possible. But why take squares?
As mentioned, squaring numbers returns a positive result. Therefore the square gives some idea of the 'absolute' value of the residuals (and doesn't allow negative residuals to cancel out positve ones). But why not just take the absolute value?
Because the square function penalises large deviations heavily, least squares prefers lines that keep every point reasonably close to the fitted value rather than letting a few points drift far away.
:::

### Fitting a model using least squares

We call the estimates $\hat{\beta_0}$ and $\hat{\beta_1}$ which minimise the sum of squares the *least squares estimates*. Some nice mathematics tells us that the least squares esimates for a simple linear model are *unique* - that is, there is one set of values for $\hat{\beta_0}$ and $\hat{\beta_1}$ which satisfy this property. Moreover, we don't have to manually adjust our parameters to keep lowering SSR - there is a convenient closed form expression which allows us to compute $\hat{\beta_0}$ and $\hat{\beta_1}$ very efficiently. As a statistical programming language, this is something R does very easily..


#### The `lm()` function

In R the `lm()` function computes the least squares estimates $\hat{\beta_0}$ and $\hat{\beta_1}$ for our simple linear model in a single command, printing both of the fitted coefficients. 

```{webr}
#| input:
#|   - linearData
lm_1 <- lm(y~x,linearData)
lm_1
```

How do these estimates compare with yours? Lets compare the lines graphically:

::: {.panel-tabset}

## plot

```{webr}
#| envir: Ex2
# we extract the coefficients from the lm object by indexing with `$`
ls_beta_0_hat <- lm_1$coefficients[1]
ls_beta_1_hat <- lm_1$coefficients[2]

ggplot(linearData)+
  aes(x=x,y=y)+
  geom_point()+
  geom_abline(
    intercept=ls_beta_0_hat,
    slope=ls_beta_1_hat)+
    geom_abline(
    intercept=beta_0_hat,
    slope=beta_1_hat)
```
:::
-   [ ] Continue

### Estimating $\sigma$

Once the line has been fitted, we estimate the spread of the error terms (i.e. $\sigma$) by measuring the spread of the residuals around our line of best fit.

We can extract the individual residual values from the fitted `lm` object:

```{webr}
e <- lm_1$residuals
ggplot(e)+
  aes(x=e)+
  geom_hist()
```
::: Exercise 
While we will often work with `lm` objects, it can be helpful to know how to calculate values ourselves.
Complete the code below to calculate the residuals for the least-squares model:
```{webr}
Y_hat <- ls_beta_0_hat + ls_beta_0_hat*linearData$x

myresiduals <- - Y_hat
```
:::

We can check that these values are correct

```{webr}
e==myresiduals
```

Dividing the residual sum of squares by $n-2$ (one degree of freedom lost per fitted coefficient) gives the mean squared error, and taking the square root yields the residual standard deviation (aka. the {{< glossary Residual-Standard-Error >}})
$$
\hat \sigma = \sqrt{\frac{1}{n-2}\sum_{i=1}^{n} r_i^2}.
$$
In R output, this value is often called the Standard error. Interpreting $\hat \sigma$ is often easier on the original $y$ scale: a typical observation falls about $\hat \sigma$ units away from the fitted line.

- [ ] Continue

## Checking assumptions

### Independence of observations
-   **What it means:** Linear regression relies on the observations being independent; each response should be unaffected by the others.
-   **Common violations:** Time series data, spatial measurements, or repeated observations on the same unit often introduce dependence.
-   **How to diagnose:** Look for clues in the study design and plot residuals in their collection order—long runs of positive or negative residuals hint at dependence.
-   **Possible remedies:** Incorporate correlation structures (e.g., mixed models), difference consecutive observations, or collect data to ensure independence.

### Linear relationships
-   **What it means:** The expected response should change linearly with the predictor.
-   **Diagnostics:** Scatterplots of $(x_i, y_i)$ and residual-versus-fitted plots are useful checks.
-   **Warning signs:** Curved patterns or systematic structure in these plots suggest the linear form is inadequate.
-   **Possible remedies:** Try transforming variables, adding polynomial terms, or introducing additional predictors.

### Constant Variance
-   **What it means:** Also called homoscedasticity, the variability of residuals should stay roughly constant across fitted values.
-   **Diagnostics:** Residual-versus-fitted plots should show an even vertical spread.
-   **Warning signs:** Funnel or fan shapes (narrow for small $\hat y$ but wide for large $\hat y$) reveal heteroscedasticity.
-   **Possible remedies:** Transform the response, use weighted least squares, or model the variance explicitly.

### Normality of error
-   **What it means:** Least squares inference (confidence intervals, hypothesis tests) assumes residuals are approximately normal.
-   **Diagnostics:** Use a normal Q-Q plot to compare residual quantiles with the theoretical normal line.
-   **Warning signs:** Systematic bends or heavy tail departures from the diagonal indicate skewness or heavy tails.
-   **Possible remedies:** Transform the response, refit with robust methods, or collect more data to stabilise the distribution.

### Examples of assumption violations
-   **Simulated example:** In this case the assumptions hold—as expected because the data were sampled from an actual linear model.
-   **Real-world contrasts:** Add case studies where independence fails (e.g., repeated measures), the relationship is curved, the variance changes, or the residuals are non-normal.
-   **Teaching tip:** For each violation, include visuals showing the problematic residual pattern and a suggested fix.

::: Example
**Independence**

Repeated measurements on the same individual violate the independence assumption because the shared subject creates dependence between observations. In the data below we record eight practice trials each week for a single violinist preparing for an audition. Within a week, longer practice hours are actually associated with slightly *lower* accuracy as fatigue sets in, but the week-level averages trend upward because the player gradually improves. The solid black regression fitted to the weekly averages therefore suggests a positive relationship that does not exist for the day-to-day observations—an ecological fallacy. Remedies such as multilevel or mixed-effects models can account for the nested structure, but they are beyond the scope of this course.

```{webr}
set.seed(2024)
independence_example <- expand_grid(
  session = factor(1:4, labels = paste("Week", 1:4)),
  trial = 1:8
) |>
  mutate(
    hours_practice = rnorm(n(), 0, 0.4) + as.integer(session),
    accuracy = 70 + 6 * as.integer(session) - 4 * (hours_practice - as.integer(session)) + rnorm(n(), 0, 1.2)
  )

session_means <- independence_example |>
  group_by(session) |>
  summarise(
    mean_hours = mean(hours_practice),
    mean_accuracy = mean(accuracy),
    .groups = "drop"
  )

ggplot(independence_example, aes(hours_practice, accuracy, colour = session)) +
  geom_point(alpha = 0.6) +
  geom_smooth(method = "lm", se = FALSE, linetype = "dashed") +
  geom_point(
    data = session_means,
    aes(mean_hours, mean_accuracy),
    inherit.aes = FALSE,
    size = 3,
    colour = "black"
  ) +
  geom_smooth(
    data = session_means,
    aes(mean_hours, mean_accuracy),
    inherit.aes = FALSE,
    method = "lm",
    se = FALSE,
    colour = "black"
  ) +
  labs(
    x = "Hours of practice",
    y = "Intonation accuracy (%)",
    colour = "Week"
  ) +
  theme(legend.position = "bottom")
```
:::

::: Example
**Linearity**

If the true relationship between the predictor and response is curved, the fitted straight line misses systematic structure. Here a quadratic mean trend leaves a clear arc in the scatterplot and in the dashed least-squares fit.

```{webr}
set.seed(2024)
nonlinear_example <- tibble(
  dosage = seq(0, 10, length.out = 60),
  response = 5 + 1.4 * dosage - 0.18 * dosage^2 + rnorm(60, 0, 1.2)
)

ggplot(nonlinear_example, aes(dosage, response)) +
  geom_point() +
  geom_smooth(method = "lm", se = FALSE, colour = "#d73027", linetype = "dashed") +
  geom_smooth(method = "loess", se = FALSE, colour = "#4575b4") +
  labs(x = "Dosage (mg)", y = "Response score") +
  theme(legend.position = "none")
```
:::

::: Example
**Constant variance**

When variability grows with the fitted value, the residuals form a funnel shape. In this simulated data the spread of the errors increases with the predictor, so the residual-versus-fitted plot fans out instead of forming an even band.

```{webr}
set.seed(2024)
hetero_example <- tibble(
  x = seq(0, 10, length.out = 80),
  y = 2 + 0.5 * x + rnorm(80, 0, 0.4 + 0.25 * x)
)

hetero_model <- lm(y ~ x, data = hetero_example)
hetero_diag <- hetero_example |>
  mutate(
    fitted = fitted(hetero_model),
    residuals = resid(hetero_model)
  )

ggplot(hetero_diag, aes(fitted, residuals)) +
  geom_hline(yintercept = 0, colour = "grey70") +
  geom_point() +
  geom_smooth(method = "loess", se = FALSE, colour = "#542788") +
  labs(x = "Fitted value", y = "Residual")
```
:::

::: Example
**Normality**

Heavy-tailed errors lead to extreme residuals that stray from the reference line on a Q-Q plot. The simulated model below uses Student-$t$ noise with only two degrees of freedom, producing thicker tails than the normal distribution.

```{webr}
set.seed(2024)
heavy_tail_example <- tibble(
  hours = runif(100, 0, 8),
  satisfaction = 4 + 0.8 * hours + rt(100, df = 2)
)

heavy_tail_model <- lm(satisfaction ~ hours, data = heavy_tail_example)
heavy_tail_resid <- tibble(residuals = resid(heavy_tail_model))

ggplot(heavy_tail_resid, aes(sample = residuals)) +
  stat_qq() +
  stat_qq_line(colour = "#1b9e77") +
  labs(x = "Theoretical quantiles", y = "Sample quantiles")
```
:::


## Inference for Simple Linear Regression


-   **Connection to simulation:** We have been given data sampled from a linear model—consider whether the parameter estimates obtained so far represent the underlying model well.
-   **Repeat sampling experiment:** Since we can simulate data, take another sample from the same model to explore variability in the estimates.
-   **Comparison visuals:** Overlay lines from multiple samples to build intuition for the sampling distribution of $\hat\beta$.
-   **Key takeaway:** Even with the same generating process, estimates vary; inference quantifies this uncertainty.

- since we have the opportunity, lets take another sample from the same model:

```{webr}
n <- 20
x <- rnorm(n,0,4)
error <-  rnorm(n,0, 2)
y <- 0.5 * x - 3 + error
linearData_new <- data.frame(x,y)
```

And fit a linear regression via least squares
```{webr}
lm_2 <- lm(y~x,linearData_new)
lm
```
lets compare out two datasets (remember, these come from the same underlying linear model with $\beta_0=-3$ and $\beta_1=0.5$ and their lines of best fit:
```{webr}
ggplot()+
  geom_point(data=linearData, aes())
```
- Repeating this process produces the sampling distribution of our parameters.

- in the real world we wont have the luxury of generating new samples by running R commands, so we have to make do with the data we have. Thankfully, if our assumptions of the underlying model are true then we can expect the same process of convergence to a stable sampling distribution to hold true and our parameter estimates for beta to have a t-distribution.

### Inference about $\beta$

-   **Objective:** Derive sampling distributions for $\hat{\beta}_0$ and $\hat{\beta}_1$ under the model assumptions.
-   **Key results:** Emphasise unbiasedness, variance formulas, and the joint normality of the estimators.
-   **Confidence intervals:** Outline the structure $\hat{\beta} \pm t_{n-2,\alpha/2} \cdot \operatorname{SE}(\hat{\beta})$.
-   **Hypothesis tests:** Summarise how to test linear contrasts of the coefficients using t-statistics.

### Inference about $\sigma$

-   **Distributional result:** With the assumptions above, the scaled residual sum of squares follows a $\chi^2_{n-2}$ distribution.
-   **Confidence interval:** Show how this leads to bounds for $\sigma$ using chi-squared quantiles.
-   **Hypothesis test:** Note the form of tests comparing error variance claims.
-   **Software link:** Point to the `sigma` output and degrees of freedom in R for practical calculation.
### Inference about $\beta_1$

-   **Hypothesis:** To assess association, test $H_0: \beta_1 = 0$ versus an appropriate alternative.
-   **Test statistic:** Introduce the t-statistic $t = \frac{\hat \beta_1 - 0}{\operatorname{SE}(\hat \beta_1)}$ and its $t_{n-2}$ reference distribution.
-   **Software output:** Interpret R's t-statistic, p-value, and confidence interval for $\beta_1$.
-   **Practical interpretation:** Translate statistical significance into statements about direction and strength of association.

## Interpreting \`summary()\` output

-   **Purpose:** `summary()` wraps the model fit, assumptions, and inference into a single report.
-   **Coefficients table:** Highlight estimates, standard errors, t-values, and p-values for $\beta_0$ and $\beta_1$.
-   **Residual standard error:** Connect this to $\hat \sigma$ and the degrees of freedom shown in the output.
-   **Model fit metrics:** Explain multiple $R^2$ and adjusted $R^2$ as measures of explained variation.
-   **Overall test:** Describe how the F-statistic in simple regression aligns with the $\beta_1$ t-test.
-   **Workflow tip:** Encourage students to read the table line by line, linking each quantity back to the modelling steps above.

Interpreting each component in context helps translate the statistical output into practical insight about the data.

## Using the model for prediction

-   **Prerequisite:** Only use the model for prediction after diagnostics suggest the assumptions hold.
-   **Point prediction:** Use `predict()` to obtain fitted values for new $x$.
-   **Interval estimates:** Emphasise the difference between confidence intervals for the mean response and prediction intervals for individual outcomes.
-   **Communicating uncertainty:** Always report the uncertainty associated with predictions.
-   **Scope of application:** Warn about extrapolating beyond the observed range of $x$ and discuss potential pitfalls.

test. 